## Q1 (4 points)
จงเขียนฟังก์ชัน `get_dict(file_name)` ที่รับ parameter เป็นชื่อไฟล์ แล้วสร้างและคืน dictionary ซึ่งมี
- key เป็นสตริงของรหัสสินค้า และ มี 
- value เป็น dictionary ที่มี 
    - key เป็นสตริงของขนาด (`S`, `M`,  `L` หรือ `XL`) และ 
    - value เป็นตัวเลขจำนวนของสินค้า
    
และแสดงผลลัพธ์เป็นข้อมูลใน dictionary ตามตัวอย่างการทำงาน

### ตัวอย่างการทำงาน: 
get_dict(**'stock1.txt'**) <br>
{4125: {'M': 4}, 5645: {'L': 7}, 7845: {'XL': 3}}

### ตัวอย่างการทำงาน: 
get_dict(**'stock2.txt'**)<br>
{7845: {'XL': 3}, 4125: {'S': 1, 'M': 4}, 5645: {'L': 7}, 9874: {'S': 3}}

### ตัวอย่างการทำงาน: 
get_dict('**stock3.txt**')<br>
{4125: {'S': 1, 'M': 4}, 5645: {'L': 7}, 9874: {'S': 8}, 9875: {'M': 8}, 7845: {'S': 10, 'XL': 3}}

<!-- grade_id: name, points: 0, grade: false, solution: true -->
```python
def get_dict(file_name):
    '''สร้าง dictionary ของรหัสและจำนวนสินค้า คืนค่าเป็น dictionary'''
    Dict = {}
    with open(file_name, 'r') as f:
        for ln in f:
            parts = ln.strip().split()
            product_code, size, number = int(parts[0]), parts[1], int(parts[2])

            if product_code not in Dict:
                Dict[product_code] = {size: number}
            else:
                if size in Dict[product_code]:
                    Dict[product_code][size] += number
                else:
                    Dict[product_code][size] = number
                    
    return Dict
```

```python
'''ตรวจสอบผลลัพธ์ของโปรแกรมด้วยการรันที่ cell นี้'''
'''ตัวอย่าง get_dict("ใส่ชื่อไฟล์.txt")'''
print(get_dict('stock1.txt'))
print(get_dict('stock2.txt'))
print(get_dict('stock3.txt'))
```

Output:
```
{4125: {'M': 4}, 5645: {'L': 7}, 7845: {'XL': 3}}
{7845: {'XL': 3}, 4125: {'S': 1, 'M': 4}, 5645: {'L': 7}, 9874: {'S': 3}}
{4125: {'S': 1, 'M': 4}, 5645: {'L': 7}, 9874: {'S': 8}, 9875: {'M': 8}, 7845: {'S': 10, 'XL': 3}}
```

<!-- grade_id: cell-498bfc73bc381903, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบการแสดงผลลัพธ์อย่างง่าย'''
assert get_dict('stock1.txt') == {4125: {'M': 4}, 5645: {'L': 7}, 7845: {'XL': 3}}
print('OK')
```

Output:
```
OK
```

<!-- grade_id: cell-2981a5915ebbc36e, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบการแสดงผลลัพธ์อย่างง่าย'''
assert get_dict('stock2.txt') == {7845: {'XL': 3}, 4125: {'S': 1, 'M': 4}, 5645: {'L': 7}, 9874: {'S': 3}}
print('OK')
```

Output:
```
OK
```

<!-- grade_id: exact, points: 1, grade: true, solution: false -->
```python
'''ตรวจสอบผลการคำนวณและการแสดงผลลัพธ์ให้เหมือนกับโจทย์ที่กำหนดให้'''
assert get_dict('stock1.txt') == {4125: {'M': 4}, 5645: {'L': 7}, 7845: {'XL': 3}}
assert get_dict('stock2.txt') == {7845: {'XL': 3}, 4125: {'S': 1, 'M': 4}, 5645: {'L': 7}, 9874: {'S': 3}}
assert get_dict('stock3.txt') == {4125: {'S': 1, 'M': 4}, 5645: {'L': 7}, 9874: {'S': 8}, 9875: {'M': 8}, 7845: {'S': 10, 'XL': 3}}
print('OK')
```

Output:
```
OK
```

<!-- grade_id: cell-c73870cae5a59dfa, points: 1, grade: true, solution: false -->
```python
'''เทสลับ1'''
```

Output:
```
'เทสลับ1'
```